# Part 1: Forward Mode Automatic Differentiation

Forward mode AD can simply be implemented by defining a class to represent [dual numbers](https://en.wikipedia.org/wiki/Dual_number) which hold the value and its derivative. The following skeleton defines a dual number and implements multiplication. 

__Tasks:__

- Addition (`__add__`) is incomplete - can you finish it? 
- Can you also implement division (`__truediv__`), subtraction (`__sub__`) and power (`__pow__`)?

<!-- grade_id: cell-a645b47036f56c2e, points: 0, grade: false, solution: true -->
```python
import math

class DualNumber:
    def __init__(self, value, dvalue):
        self.value = value
        self.dvalue = dvalue

    def __str__(self):
        return str(self.value) + " + " + str(self.dvalue) + "ε"

    def __mul__(self, other):
        return DualNumber(self.value * other.value,
            self.dvalue * other.value + other.dvalue * self.value)
    
    def __add__(self, other):
        #TODO: finish me
        # YOUR CODE HERE
        return DualNumber(self.value + other.value, self.dvalue + other.dvalue)
        #raise NotImplementedError()
    
    def __sub__(self, other):
        #TODO: finish me
        # YOUR CODE HERE
        return DualNumber(self.value - other.value, self.dvalue - other.dvalue)
      
    def __truediv__(self, other):
      return DualNumber(self.value / other.value, (self.dvalue/other.value - self.value * other.dvalue/other.value ** 2))
      #return DualNumber(self.value / other.value, (self.value*other.dvalue - self.dvalue * other.value) / (other.value ** 2))
      
    def __pow__(self, other):
        return DualNumber(self.value ** other.value, other.value * self.value**(other.value-1) * self.dvalue + math.log(self.value) * self.value ** other.value * other.dvalue)
      
      
    # TODO: add missing methods
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-7923d5c39d574cd7, points: 3, grade: true, solution: false -->
```python
# Tests

print(DualNumber(1,0) + DualNumber(1,0) / DualNumber(1,0) - DualNumber(1,0)**DualNumber(1,0))

```

Output:
```
1.0 + 0.0ε
```

## Implementing math functions

We also need to implement some core math functions. Here's the sine function for a dual number:

```python
def sin(x):
    return DualNumber(math.sin(x.value), math.cos(x.value)*x.dvalue)
```

__Task:__ can you implement the _cosine_ (`cos`), _tangent_ (`tan`), and _exponential_ (`exp`) functions in the code block below?

<!-- grade_id: cell-61b9462cb781a8e1, points: 0, grade: false, solution: true -->
```python
# TODO: implement additional math functions on dual numbers

def cos(x):
    # YOUR CODE HERE
    return DualNumber(math.cos(x.value), -1 * math.sin(x.value) * x.dvalue)
    #raise NotImplementedError()

def tan(x):
    # YOUR CODE HERE
    return DualNumber(math.tan(x.value), (1/(math.cos(x.value) **2)) * x.dvalue)
    #raise NotImplementedError()

def exp(x):
    # YOUR CODE HERE
    return DualNumber(math.exp(x.value), math.exp(x.value) * x.dvalue)
    #raise NotImplementedError()
```

<!-- grade_id: cell-9055ec23e6e9ffb8, points: 3, grade: true, solution: false -->
```python
# Tests
assert cos(DualNumber(0,0)).value == 1
assert tan(DualNumber(0,0)).value == 0
assert exp(DualNumber(0,0)).value == 1

```

## Time to try it out

We're now in a position to try our implementation.

__Task:__ 

- Try running the following code to compute the value of the function $z=x\cdot y+sin(x)$ given $x=0.5$ and $y=4.2$, together with the derivative $\partial z/\partial x$ at that point. 

<!-- grade_id: cell-12f7c8fa42784946, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x = 0.5
dxdx = 1
dxdy = 0

###
y = 4.2
dydy = 1
dydx = 0

z = x * y + math.sin(x)

dz_dx = DualNumber(x,1) * DualNumber(y,0) + sin(DualNumber(x,1))
print(dz_dx)
#raise NotImplementedError()
```

Output:
```
2.579425538604203 + 5.077582561890373ε
```

__Task__: Differentiate the above function with respect to $x$ and write the symbolic derivatives in the following box. Verify the result computed above is correct by plugging-in the values into your symbolic gradient expression.

<!-- grade_id: cell-1ae4d0a0801a142a, points: 2, grade: true, solution: true -->
dz/dx =  y + cos(x) = 4.2 + cos(0.5) =  5.077

__Task:__ Now use the code block below to compute the derivative $\partial z/\partial y$ of the above expression (at the same point $x=0.5, y=4.2$ as above) and store the derivative in the variable `dzdy` (just the derivative, not the Dual Number). Verify by hand that the result is correct.

<!-- grade_id: cell-eab763bf1b218dc8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()

dzdy = DualNumber(x,0) * DualNumber(y,1) 

print(dzdy)

```

Output:
```
2.1 + 0.5ε
```

<!-- grade_id: cell-a2f9f0ca4b1ed763, points: 2, grade: true, solution: false -->
```python
#Tests
# z = x * y + math.sin(x)

# dzdy = x 
assert dzdy.dvalue == 0.5
 

```

__Task:__ Finally, use the code block below to experiment and test the other math functions and methods you created.

<!-- grade_id: cell-b71d4bb64928ee38, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
r = 5

s = 2

dz2_dr =   DualNumber(r,1) / DualNumber(s,0) +tan(DualNumber(r,1))

print(dz2_dr)

#raise NotImplementedError()


#dz2_dr = 1/s + 1/(cosr*cosr))
#       = 12.927
 
```

Output:
```
-0.8805150062465859 + 12.927881707458354ε
```

```python
dz3_dr = DualNumber(r,1) ** DualNumber(s,0) +  cos(DualNumber(r,1))
print(dz3_dr)

#dz3_dr =  ( s*r**s-1 ) -sin(r) = 10.95

```

Output:
```
25.283662185463225 + 10.958924274663138ε
```

```python
dz4_dr = DualNumber(r,1) ** DualNumber(s,0) *  cos(DualNumber(r,1))
print(dz4_dr)

```

Output:
```
7.091554636580656 + 26.809728721210725ε
```

```python
# value calculated by hand and value calculated using Forward AD are the same
```